
# Lab 4 Validation Notebook

You cannot write any code in this notebook – to actually work on the lab, open `lab.py` and `lab.ipynb`.

The only purpose of this notebook is to give you a blank copy of `lab.ipynb`,
so that you can go to **Kernel > Restart & Run All** and ensure that all public `grader.check` cells pass using just the code in your `lab.py`.

**Before submitting Lab 4, make sure that the call to `grader.check_all()` at the bottom of this notebook shows that all test cases passed!** 
If it doesn't, there's likely a function in `lab.ipynb` that is not implemented correctly in `lab.py`, or it could be that a function in `lab.py` depends on an object (e.g. a DataFrame) that is not an argument to that function.
    

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab.ipynb")
```

```python
%load_ext autoreload
%autoreload 2
```

```python
from lab import *
```

```python
import pandas as pd
import numpy as np
import io
from pathlib import Path
import os
```

```python
# don't change this cell -- it is needed for the tests to work
fp = Path('data')/'login_table.csv'
login = pd.read_csv(fp)
q1_result = prime_time_logins(login)
login
```

```python
grader.check("q1")
```

```python
# don't change this cell -- it is needed for the tests to work
fp = Path('data')/'login_table.csv'
login = pd.read_csv(fp)
q2_result = count_frequency(login)
```

```python
grader.check("q2")
```

```python
grader.check("q3")
```

```python
grader.check("q4")
```

```python
# don't change this cell -- it is needed for the tests to work
superheroes_fp = Path('data')/'superheroes.csv'
heroes = pd.read_csv(superheroes_fp, index_col=0)
bhbe_out = bhbe_col(heroes)

obs_stat_out = superheroes_observed_stat(heroes)

simulate_bhbe_out = simulate_bhbe_null(10)

pval_out = superheroes_calc_pval()
```

```python
grader.check("q5")
```

```python
skittles_fp = Path('data')/'skittles.tsv'
skittles = pd.read_csv(skittles_fp, sep='\t')
skittles.head()
```

```python
skittles.shape
```

```python
# don't change this cell -- it is needed for the tests to work
# cell may take about 1-2 minutes to execute to completion
skittles_fp = Path('data')/'skittles.tsv'
skittles = pd.read_csv(skittles_fp, sep='\\t', engine='python')
q6_diff_of_means_out = diff_of_means(skittles)
q6_simulate_null_out = simulate_null(skittles)
q6_pval_out = pval_color(skittles)
```

```python
grader.check("q6")
```

```python
# don't change this cell -- it is needed for the tests to work
q7_out = ordered_colors()
q7_colors = {'green', 'orange', 'purple', 'red', 'yellow'}
q7_test_colors = [x[0] for x in q7_out]
```

```python
grader.check("q7")
```

```python
# don't change this cell -- it is needed for the tests to work
q8_out = same_color_distribution()
```

```python
grader.check("q8")
```

```python
# don't change this cell -- it is needed for the tests to work
q9_out = perm_vs_hyp()
```

```python
grader.check("q9")
```

```python
grader.check_all()
```


If you were able to go to **Kernel > Restart & Run All** and see all test cases pass above, and you've thoroughly tested your code yourself, you're ready to submit `lab.py` to Gradescope!
    

Tests for q1 (points: None)
```
>>> len(q1_result) == 433
True
```
```
>>> q1_result.loc[393, "Time"] > 9
True
```
```
>>> q1_result.loc[457, 'Time'] == 4
True
```
failure_message: check user 457
```
>>> q1_result.loc[458, 'Time'] > 50
True
```
failure_message: check user 457
```
>>> len(set(q1_result.index)) == 433
True
```
failure_message: check number of unique users


Tests for q2 (points: None)
```
>>> len(q2_result) == 433
True
```
```
>>> np.isclose(q2_result.loc[466], 0.24250681198910082)
True
```


Tests for q3 (points: None)
```
>>> isinstance(suits_null_hyp(), list)
True
```
```
>>> set(suits_null_hyp()).issubset({1, 2, 3, 4})
True
```
```
>>> pd.Series(simulate_suits_null()).isin([0, 1]).all()
True
```
```
>>> 0 < estimate_suits_p_val(1000) < 0.1
True
```


Tests for q4 (points: None)
```
>>> set(car_null_hypoth()) <= set(range(1, 7))
True
```
```
>>> set(car_alt_hypoth()) <= set(range(1, 7))
True
```
```
>>> set(car_test_stat()) <= set(range(1, 5))
True
```
```
>>> car_p_value() in set(range(1, 6))
True
```


Tests for q5 (points: None)
```
>>> isinstance(bhbe_out, pd.Series)
True
```
```
>>> bhbe_out.dtype == np.dtype('bool')
True
```
```
>>> bhbe_out.sum() == 93
True
```
```
>>> 0.5 <= obs_stat_out <= 1.0
True
```
```
>>> isinstance(simulate_bhbe_out, np.ndarray)
True
```
```
>>> simulate_bhbe_out.shape[0] == 10
True
```
```
>>> ((0.45 <= simulate_bhbe_out) & (simulate_bhbe_out <= 1)).all()
True
```
```
>>> len(pval_out) == 2
True
```
```
>>> 0 <= pval_out[0] <= 1
True
```
```
>>> pval_out[1] in ['Reject', 'Fail to reject']
True
```


Tests for q6 (points: None)
```
>>> isinstance(q6_diff_of_means_out, float)
True
```
```
>>> isinstance(q6_simulate_null_out, float)
True
```
```
>>> 0 <= q6_simulate_null_out <= 1.0
True
```
```
>>> isinstance(q6_pval_out, float)
True
```
```
>>> 0 <= q6_pval_out <= 1
True
```
```
>>> q6_diff_of_means_out > 0
True
```
failure_message: should be greater than zero


Tests for q7 (points: None)
```
>>> len(q7_out) == 5
True
```
```
>>> set([x[0] for x in q7_out]) == q7_colors
True
```
```
>>> all([isinstance(x[1], float) for x in q7_out])
True
```
```
>>> q7_test_colors.index('green') > q7_test_colors.index('yellow')
True
```
failure_message: yellow less than green
```
>>> np.isclose(q7_out[0][1], 0.0)
True
```
failure_message: smallest pval


Tests for q8 (points: None)
```
>>> isinstance(q8_out, tuple)
True
```
failure_message: tuple
```
>>> isinstance(q8_out[0], float)
True
```
failure_message: wrong output type at index 0
```
>>> q8_out[1] in ['Fail to Reject', 'Reject']
True
```
failure_message: wrong output type at index 1
```
>>> np.isclose(q8_out[0], 0.005, atol=0.25)
True
```
failure_message: p-value, approximate within 0.25
```
>>> np.isclose(q8_out[0], 0.005, atol=0.5)
True
```
failure_message: p-value, approximate within 0.5


Tests for q9 (points: None)
```
>>> len(q9_out) == 5
True
```
failure_message: output length should be 5
```
>>> set(q9_out) <= set(['P', 'H'])
True
```
failure_message: output contains answers other than P or H
