
# Lab 6 Validation Notebook

You cannot write any code in this notebook – to actually work on the lab, open `lab.py` and `lab.ipynb`.

The only purpose of this notebook is to give you a blank copy of `lab.ipynb`,
so that you can go to **Kernel > Restart & Run All** and ensure that all public `grader.check` cells pass using just the code in your `lab.py`.

**Before submitting Lab 6, make sure that the call to `grader.check_all()` at the bottom of this notebook shows that all test cases passed!** 
If it doesn't, there's likely a function in `lab.ipynb` that is not implemented correctly in `lab.py`, or it could be that a function in `lab.py` depends on an object (e.g. a DataFrame) that is not an argument to that function.
    

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab.ipynb")
```

```python
%load_ext autoreload
%autoreload 2
```

```python
from lab import *
```

```python
import os
import pandas as pd
import numpy as np
import requests
import bs4
import lxml
```

```python
# Don't delete this cell!
question1()
```

```python
grader.check("q1")
```

```python
# don't delete this cell, but do run it -- it is needed for the autograder tests

# public test for extract_book_links 
extract_book_links_fp = os.path.join('data', 'products.html')
extract_book_out = extract_book_links(
    open(extract_book_links_fp, encoding='utf-8').read()
)
extract_book_url = 'scarlet-the-lunar-chronicles-2_218/index.html'

# doc tests for get product info
get_product_info_fp = os.path.join('data', 'Frankenstein.html')
get_product_info_out = get_product_info(
    open(get_product_info_fp, encoding='utf-8').read(), ['Default']
)

# public test for scrape books 
scrape_books_out = scrape_books(1, ['Mystery'])
```

```python
grader.check("q2")
```

```python
# don't delete this cell, but do run it -- it is needed for the autograder tests

# public test for stock_history
history = stock_history('BYND', 2019, 6)

# public test for stock_stats
stats = stock_stats(history)
```

```python
grader.check("q3")
```

```python
# don't delete this cell, but do run it -- it is needed for the autograder tests
comments = get_comments(18344932)
```

```python
grader.check("q4")
```

```python
grader.check_all()
```


If you were able to go to **Kernel > Restart & Run All** and see all test cases pass above, and you've thoroughly tested your code yourself, you're ready to submit `lab.py` to Gradescope!
    

Tests for q1 (points: None)
```
>>> os.path.exists('lab06_1.html')
True
```
```
>>> q01_html = bs4.BeautifulSoup(open("lab06_1.html", encoding='utf-8'))
>>> len(q01_html.find_all('image')) == 0
True
```
failure_message: you should not use the <image> tag to add images


Tests for q2 (points: None)
```
>>> extract_book_out[1] == extract_book_url
True
```
```
>>> isinstance(get_product_info_out, dict)
True
```
```
>>> 'Category' in get_product_info_out.keys()
True
```
```
>>> get_product_info_out['Rating'] == 'Two'
True
```
```
>>> scrape_books_out.shape == (1, 11)
True
```
```
>>> scrape_books_out['Rating'][0] == 'Four'
True
```
```
>>> scrape_books_out['Title'][0] == 'Sharp Objects'
True
```


Tests for q3 (points: None)
```
>>> history.shape == (20, 13)
True
```
```
>>> history.label.iloc[-1] == 'June 03, 19'
True
```
```
>>> (len(stats[0]), len(stats[1])) == (7, 6)
True
```
```
>>> (float(stats[0][1:-1]) > 30) == True
True
```
```
>>> float(stats[1][:-1]) > 1 == True
True
```
```
>>> stats[1][-1] == 'B'
True
```


Tests for q4 (points: None)
```
>>> comments.shape == (18, 5)
True
```
```
>>> comments.loc[5, 'by'] == 'RobAtticus'
True
```
```
>>> comments.loc[5, 'time'].day == 31
True
```
