
# Project 1 Validation Notebook

You cannot write any code in this notebook – to actually work on the project, open `project.py` and `project.ipynb`.

The only purpose of this notebook is to give you a blank copy of `project.ipynb`,
so that you can go to **Kernel > Restart & Run All** and ensure that all public `grader.check` cells pass using just the code in your `project.py`.

**Before submitting Project 1, make sure that the call to `grader.check_all()` at the bottom of this notebook shows that all test cases passed!** 
If it doesn't, there's likely a function in `project.ipynb` that is not implemented correctly in `project.py`, or it could be that a function in `project.py` depends on an object (e.g. a DataFrame) that is not an argument to that function.
    

```python
# Initialize Otter
import otter
grader = otter.Notebook("project.ipynb")
```

```python
%load_ext autoreload
%autoreload 2
```

```python
import pandas as pd
import numpy as np
import os
```

```python
from project import *
```

```python
grades_fp = os.path.join('data', 'grades.csv')
grades = pd.read_csv(grades_fp)
grades.head()
```

```python
grader.check("q1")
```

```python
grader.check("q2")
```

```python
grader.check("q3")
```

```python
grader.check("q4")
```

```python
grader.check("q5")
```

```python
grader.check("q6")
```

```python
grader.check("q7")
```

```python
grader.check("q8")
```

```python
final_breakdown_fp = os.path.join('data', 'final_exam_breakdown.csv')
final_breakdown = pd.read_csv(final_breakdown_fp)
final_breakdown.head()
```

```python
final_breakdown.shape
```

```python
grades.loc[grades['PID'] == 'A99381181', 'Final']
```

```python
grader.check("q9")
```

```python
grades_combined = combine_grades(grades, raw_redemption(final_breakdown, [1, 2, 3, 7, 9, 12]))
grades_combined.head()
```

```python
grader.check("q10")
```

```python
grader.check("q11")
```

```python
grader.check_all()
```


If you were able to go to **Kernel > Restart & Run All** and see all test cases pass above, and you've thoroughly tested your code yourself, you're ready to submit `project.py` to Gradescope!
    

Tests for q1 (points: None)
```
>>> out = get_assignment_names(grades)
>>> out['final'] == ['Final']
True
```
```
>>> out = get_assignment_names(grades)
>>> set(out.keys()) == {'lab', 'project', 'midterm', 'final', 'disc', 'checkpoint'}
True
```
```
>>> out = get_assignment_names(grades)
>>> 'project02' in out['project']
True
```


Tests for q10 (points: None)
```
>>> out = z_score(grades['Final'])
>>> out.dtype == float and out.shape[0] == grades.shape[0]
True
```
```
>>> out = add_post_redemption(grades_combined)
>>> 'Midterm Score Pre-Redemption' in out and 'Midterm Score Post-Redemption' in out
True
```
```
>>> out = add_post_redemption(grades_combined)
>>> np.any(np.isclose(out['Midterm Score Pre-Redemption'], out['Midterm Score Post-Redemption']))
True
```
failure_message: Not all students' midterm scores should change!


Tests for q11 (points: None)
```
>>> total_points_post_redemption(grades_combined).dtypes == float
True
```
```
>>> isinstance(proportion_improved(grades_combined), float)
True
```
```
>>> 0.07 < proportion_improved(grades_combined) < 0.12
True
```
failure_message: Check proportion_improved with public grades_combined


Tests for q2 (points: None)
```
>>> out = projects_total(grades)
>>> np.all((0 <= out) & (out <= 1))
True
```
```
>>> out = projects_total(grades)
>>> 0.7 < out.mean() < 0.9
True
```


Tests for q3 (points: None)
```
>>> out = last_minute_submissions(grades)
>>> isinstance(out, pd.Series)
True
```
```
>>> out = last_minute_submissions(grades)
>>> np.all(out.index == ['lab0%d' % d for d in range(1, 10)])
True
```
failure_message: Checking that all labs, and only labs, are included in the results of last_minute_submissions.
```
>>> out = last_minute_submissions(grades)
>>> (out > 0).sum() == 8
True
```
failure_message: Checking that there are exactly 8 labs that are incorrectly marked as late.


Tests for q4 (points: None)
```
>>> out = lateness_penalty(grades['lab01 - Lateness (H:M:S)'])
>>> isinstance(out, pd.Series)
True
```
```
>>> out = lateness_penalty(grades['lab01 - Lateness (H:M:S)'])
>>> set(out.unique()) <= {1.0, 0.9, 0.7, 0.4}
True
```


Tests for q5 (points: None)
```
>>> out = process_labs(grades)
>>> out.columns.tolist() == ['lab%02d' % x for x in range(1, 10)]
True
```
```
>>> out = process_labs(grades)
>>> np.all((0.60 <= out.mean()) & (out.mean() <= 0.90))
True
```


Tests for q6 (points: None)
```
>>> cols = 'lab01 lab02 lab03'.split()
>>> processed = pd.DataFrame([[0.2, 0.90, 1.0]], index=[0], columns=cols)
>>> np.isclose(lab_total(processed), 0.95).all()
True
```


Tests for q7 (points: None)
```
>>> out = total_points(grades)
>>> np.all((0 <= out) & (out <= 1))
True
```
```
>>> out = total_points(grades)
>>> 0.7 < out.mean() < 0.9
True
```


Tests for q8 (points: None)
```
>>> out = final_grades(pd.Series([0.92, 0.81, 0.41]))
>>> np.all(out == ['A', 'B', 'F'])
True
```
```
>>> total = total_points(grades)
>>> out = letter_proportions(total)
>>> np.all(out.index == ['B', 'C', 'A', 'D', 'F'])
True
```
```
>>> total = total_points(grades)
>>> out = letter_proportions(total)
>>> out.sum() == 1.0
True
```


Tests for q9 (points: None)
```
>>> out = raw_redemption(final_breakdown, [1])
>>> 'PID' in out.columns and 'Raw Redemption Score' in out.columns
True
```
failure_message: Check the column names of the DataFrame returned by raw_redemption!
```
>>> out = combine_grades(grades, raw_redemption(final_breakdown, [1]))
>>> out.shape == (535, 101) and 'Raw Redemption Score' in out.columns
True
```
failure_message: Check the shape of the DataFrame returned by combine_grades!
