
# Lab 1 Validation Notebook

You cannot write any code in this notebook – to actually work on the lab, open `lab.py` and `lab.ipynb`.

The only purpose of this notebook is to give you a blank copy of `lab.ipynb`,
so that you can go to **Kernel > Restart & Run All** and ensure that all public `grader.check` cells pass using just the code in your `lab.py`.

**Before submitting Lab 1, make sure that the call to `grader.check_all()` at the bottom of this notebook shows that all test cases passed!** 
If it doesn't, there's likely a function in `lab.ipynb` that is not implemented correctly in `lab.py`, or it could be that a function in `lab.py` depends on an object (e.g. a DataFrame) that is not an argument to that function.
    

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab.ipynb")
```

```python
%load_ext autoreload
%autoreload 2
```

```python
from lab import *
```

```python
import os
import io
import pandas as pd
import numpy as np
```

```python
from IPython.display import YouTubeVideo
YouTubeVideo('PPKXJqu2XmY')
```

```python
# The cells below are here for you to write scratch work in. 
# You should write the code for your answer in `lab.py`, not here.
```

```python
grader.check("q0")
```

```python
grader.check("q1")
```

```python
%%time
same_diff_ints([5, 3, 1, 5, 9, 8])
```

```python
grader.check("q2")
```

```python
grader.check("q3")
```

```python
grader.check("q4")
```

```python
fp = os.path.join('data', 'chars.txt')
last_chars(open(fp))
```

```python
grader.check("q5")
```

```python
# Don't change this cell -- it is needed for the tests to work
A_1 = np.array([2, 4, 6, 7])
out_1 = add_root(A_1)

A_2 = np.array([1, 2, 16, 17, 32, 49])
out_2 = where_square(A_2)
```

```python
grader.check("q6")
```

```python
# Don't change this cell -- it is needed for the tests to work
fp = os.path.join('data', 'stocks.csv')
stocks = np.array([float(x) for x in open(fp)])
out_3_stocks = growth_rates(stocks)

A_4 = np.array([3, 3, 3, 3])
out_4 = with_leftover(A_4)
```

```python
grader.check("q7")
```

```python
# Do not edit this cell -- it is needed for the tests
salary_fp = os.path.join('data', 'salary.csv')
salary = pd.read_csv(salary_fp)
salary.head()
```

```python
# Do not edit this cell -- it is needed for the tests
salary_fp = os.path.join('data', 'salary.csv')
salary = pd.read_csv(salary_fp)
stats = salary_stats(salary)

salary_sample = pd.read_csv('data/salary_sample.csv')
sample_stats = salary_stats(salary_sample)
```

```python
grader.check("q8")
```

```python
# Do not edit -- needed for tests
fp = os.path.join('data', 'malformed.csv')
cols = ['first', 'last', 'weight', 'height', 'geo']
df = parse_malformed(fp)
dg = pd.read_csv(fp, nrows=4, skiprows=10, names=cols)
```

```python
grader.check("q9")
```

```python
grader.check_all()
```


If you were able to go to **Kernel > Restart & Run All** and see all test cases pass above, and you've thoroughly tested your code yourself, you're ready to submit `lab.py` to Gradescope!
    

Tests for q0 (points: 0)
```
>>> consecutive_ints([5, 3, 6, 4, 9, 8])
True
```
```
>>> consecutive_ints([1, 3, 5, 7, 9]) == False
True
```
```
>>> consecutive_ints([]) == False
True
```


Tests for q1 (points: None)
```
>>> median_vs_mean([6, 5, 4, 3, 2])
True
```
```
>>> median_vs_mean([50, 20, 15, 40])
True
```
```
>>> median_vs_mean([1, 8, 9]) == False
True
```


Tests for q2 (points: None)
```
>>> same_diff_ints([5, 3, 1, 5, 9, 8])
True
```
```
>>> same_diff_ints([1, 3, 5, 7, 9]) == False
True
```


Tests for q3 (points: None)
```
>>> n_prefixes('Billy', 4) == 'BillBilBiB'
True
```
```
>>> n_prefixes('Marina', 3) == 'MarMaM'
True
```
```
>>> n_prefixes('aaron', 2) == 'aaa'
True
```
```
>>> n_prefixes('Justin', 5) == 'JustiJustJusJuJ'
True
```


Tests for q4 (points: None)
```
>>> exploded_numbers([3, 8, 15], 2) == ['01 02 03 04 05', '06 07 08 09 10', '13 14 15 16 17']
True
```
```
>>> exploded_numbers([9, 99], 3) == ['006 007 008 009 010 011 012', '096 097 098 099 100 101 102']
True
```


Tests for q5 (points: None)
```
>>> fp = os.path.join('data', 'chars.txt')
>>> last_chars(open(fp)) == 'hrg'
True
```


Tests for q6 (points: None)
```
>>> isinstance(out_1, np.ndarray)
True
```
```
>>> np.all(out_1 >= A_1)
True
```
```
>>> np.isclose(out_1[3], 7 + np.sqrt(3))
True
```
```
>>> isinstance(out_2, np.ndarray)
True
```
```
>>> out_2.dtype == np.dtype('bool')
True
```
```
>>> out_2[2]
True
```


Tests for q7 (points: None)
```
>>> out_3_stocks.dtype == np.dtype('float')
True
```
```
>>> out_3_stocks.max() == 0.03
True
```
```
>>> import numbers
>>> isinstance(out_4, numbers.Integral)
True
```
```
>>> out_4 == 1
True
```


Tests for q8 (points: None)
```
>>> isinstance(stats, pd.Series)
True
```
```
>>> 'total_highest' in stats.index
True
```
```
>>> isinstance(stats.loc['duplicates'], bool) or isinstance(stats.loc['duplicates'], np.bool_)
True
```


Tests for q9 (points: None)
```
>>> list(df.columns) == cols
True
```
```
>>> df['last'].dtype == np.dtype('O')
True
```
```
>>> df['height'].dtype == np.dtype('float64')
True
```
```
>>> df['geo'].str.contains(',').all()
True
```
```
>>> len(df) == 100
True
```
```
>>> dg.index = range(9, 13)
>>> (dg == df.iloc[9:13]).all().all()
True
```
failure_message: doctest examples
