
# Project 2 Validation Notebook

You cannot write any code in this notebook – to actually work on the project, open `project.py` and `project.ipynb`.

The only purpose of this notebook is to give you a blank copy of `project.ipynb`,
so that you can go to **Kernel > Restart & Run All** and ensure that all public `grader.check` cells pass using just the code in your `project.py`.

**Before submitting Project 2, make sure that the call to `grader.check_all()` at the bottom of this notebook shows that all test cases passed!** 
If it doesn't, there's likely a function in `project.ipynb` that is not implemented correctly in `project.py`, or it could be that a function in `project.py` depends on an object (e.g. a DataFrame) that is not an argument to that function.
    

```python
# Initialize Otter
import otter
grader = otter.Notebook("project.ipynb")
```

```python
%load_ext autoreload
%autoreload 2
```

```python
import numpy as np
import pandas as pd
import pathlib
from pathlib import Path

from IPython.display import clear_output

import urllib.request
```

```python
from project import *
```

```python
from ipywidgets import widgets, interact, FloatSlider
import plotly.express as px
```

```python
# data_dir will be use for multiple cell
data_dir = Path('data')
vacs_path = data_dir / 'covid_vaccinations.csv'
vacs = pd.read_csv(vacs_path)
vacs
```

```python
DEFAULT = 'US'

def plot_cases(country):
    country_only = vacs[vacs['Country_Region'] == country]
    fig = px.line(country_only, x='Date', y='Doses_admin', title=f"'Doses_admin' column for {country}")
    fig.show()
    

dropdown_cases = widgets.Dropdown(options=np.sort(vacs['Country_Region'].unique()), value=DEFAULT)

def dropdown_cases_handler(change):
    if change['name'] == 'value' and (change['new'] != change['old']):
        clear_output()
        display(dropdown_cases)
        plot_cases(change['new'])
        
display(dropdown_cases)
plot_cases(DEFAULT)
dropdown_cases.observe(dropdown_cases_handler)
```

```python
# run this to see the result, and don't change this cell -- it is used by the tests
mono = monotonic_violations_by_country(vacs)
mono
```

```python
grader.check("q1")
```

```python
# run this to see the result, and don't change this cell -- it is used by the tests
tots = robust_totals(vacs)
tots
```

```python
grader.check("q2")
```

```python
pops_raw = pd.read_csv(data_dir / 'populations.csv')
pops_raw
```

```python
# run this to see the result, and don't change this cell -- it is used by the tests
pops = fix_dtypes(pops_raw)
pops
```

```python
grader.check("q3")
```

```python
# run this to see the result, and don't change this cell -- it is used by the tests
missing = missing_in_pops(tots, pops)
missing
```

```python
# run this to see the result, and don't change this cell -- it is used by the tests
pops_fixed = fix_names(pops)
pops_fixed
```

```python
grader.check("q4")
```

```python
# run this to see the result, and don't change this cell -- it is used by the tests
fig = draw_choropleth(tots, pops_fixed)
fig.show()
```

```python
grader.check("q5")
```

```python
# Download Israeli COVID vaccinations data from the ☁️
if not pathlib.Path(data_dir / 'israel.csv').exists():
    urllib.request.urlretrieve(
        'https://f000.backblazeb2.com/file/dsc-data/covid-israel/israel.csv',
        data_dir / 'israel.csv'
    )
```

```python
israel_raw = pd.read_csv(data_dir / 'israel.csv')
israel_raw
```

```python
# run this to see the result, and don't change this cell -- it is used by the tests
israel = clean_israel_data(israel_raw)
israel
```

```python
grader.check("q6")
```

```python
israel_raw.info()
```

```python
israel.info()
```

```python
israel.assign(null_age=israel['Age'].isna()).groupby('null_age').mean()
```

```python
# run this to see the result, and don't change this cell -- it is used by the tests
# expect this cell to take a while to run!
test_stats_vax, test_stats_sick = mcar_permutation_tests(israel, 100)
```

```python
grader.check("q7")
```

```python
vax = israel.dropna()
```

```python
# run this to see the result, and don't change this cell -- it is used by the tests
eff_overall = effectiveness(vax)
eff_overall
```

```python
grader.check("q8")
```

```python
AGE_GROUPS = [
    '12-15',
    '16-19',
    '20-29',
    '30-39',
    '40-49',
    '50-59',
    '60-69',
    '70-79',
    '80-89',
    '90-'
]
```

```python
# run this to see the result, and don't change this cell -- it is used by the tests
s_eff = stratified_effectiveness(vax)
s_eff
```

```python
grader.check("q9")
```

```python
# run this to see the result, and don't change this cell -- it is used by the tests
eff_example = effectiveness_calculator(
    young_vaccinated_prop=0.01,
    old_vaccinated_prop=0.99,
    young_risk_vaccinated=0.01,
    young_risk_unvaccinated=0.20,
    old_risk_vaccinated=0.10,
    old_risk_unvaccinated=0.50
)
eff_example
```

```python
grader.check("q10")
```

```python
# The sliders below are in the same order as the arguments below,
# in case some of the slider labels are cut off.
interact(effectiveness_calculator,
                 young_vaccinated_prop=FloatSlider(min=0.01, max=0.99, step=0.01, value=0.01),
                 old_vaccinated_prop=FloatSlider(min=0.01, max=0.99, step=0.01, value=0.99),
                 young_risk_vaccinated=FloatSlider(min=0.01, max=0.99, step=0.01, value=0.01),
                 young_risk_unvaccinated=FloatSlider(min=0.01, max=0.99, step=0.01, value=0.20),
                 old_risk_vaccinated=FloatSlider(min=0.01, max=0.99, step=0.01, value=0.10),
                 old_risk_unvaccinated=FloatSlider(min=0.01, max=0.99, step=0.01, value=0.50));
```

```python
# run this to see the result, and don't change this cell -- it is used by the tests
eff_extreme = effectiveness_calculator(**extreme_example())
eff_extreme
```

```python
grader.check("q11")
```

```python
grader.check_all()
```


If you were able to go to **Kernel > Restart & Run All** and see all test cases pass above, and you've thoroughly tested your code yourself, you're ready to submit `project.py` to Gradescope!
    

Tests for q1 (points: None)
```
>>> count_monotonic(np.array([3, 6, 6, 2, 5, 8])) == 1
True
```
```
>>> isinstance(mono, pd.DataFrame)
True
```
```
>>> mono.shape == (194, 2) and set(mono.columns) == {'Doses_admin_monotonic', 'People_at_least_one_dose_monotonic'}
True
```
```
>>> mono.sum().sum() > 0
True
```


Tests for q10 (points: None)
```
>>> isinstance(eff_example, dict)
True
```
```
>>> eff_example.keys() == {'Overall', 'Young', 'Old'}
True
```
```
>>> test_eff = effectiveness_calculator(young_vaccinated_prop=0.5, old_vaccinated_prop=0.5, young_risk_vaccinated=0.01, young_risk_unvaccinated=0.20, old_risk_vaccinated=0.01, old_risk_unvaccinated=0.20)
>>> test_eff['Overall'] == test_eff['Young'] == test_eff['Old'] == 0.95
True
```
```
>>>
```


Tests for q11 (points: None)
```
>>> isinstance(eff_extreme, dict)
True
```
```
>>> keys = {'old_risk_unvaccinated', 'young_risk_vaccinated', 'young_vaccinated_prop', 'old_risk_vaccinated', 'old_vaccinated_prop', 'young_risk_unvaccinated'}
>>> extreme_example().keys() == keys
True
```
```
>>> eff_extreme.keys() == {'Overall', 'Young', 'Old'}
True
```


Tests for q2 (points: None)
```
>>> isinstance(tots, pd.DataFrame)
True
```
```
>>> tots.shape == (194, 2)
True
```
```
>>> len(set(tots.columns).difference({'Doses_admin', 'People_at_least_one_dose'}))==0
True
```
```
>>> list(tots.dtypes) == [float]*2
True
```


Tests for q3 (points: None)
```
>>> isinstance(pops, pd.DataFrame) and pops.shape == (205, 6)
True
```
```
>>> # pops.loc[pops['Country (or dependency)'] == 'Japan', 'Population in 2023'].iloc[0] == 123294513
>>> 
>>> len(set(pops.columns).difference({'Area (Km²)',
...  'Country (or dependency)',
...  'Density (P/Km²)',
...  'ISO',
...  'Population in 2023',
...  'World Percentage'})) == 0
True
```


Tests for q4 (points: None)
```
>>> isinstance(missing, set)
True
```
```
>>> ex_tots = pd.DataFrame({
...             'Doses_admin': [1, 2, 3],
...             'People_partially_vaccinated': [1, 2, 3],
...             'People_fully_vaccinated': [1, 2, 3]
...         },
...         index = ['China', 'Angola', 'Republic of Data Science']
...     )
>>> missing_in_pops(ex_tots, pops) == {'Republic of Data Science'}
True
```
```
>>> isinstance(pops_fixed, pd.DataFrame)
True
```
```
>>> pops_fixed.shape == (205, 6)
True
```
```
>>> 'Burma' in pops_fixed['Country (or dependency)'].values
True
```
```
>>> not 'Myanmar' in pops_fixed['Country (or dependency)'].values
True
```


Tests for q5 (points: None)
```
>>> str(type(fig)) == "<class 'plotly.graph_objs._figure.Figure'>" and fig.data[0].type == 'choropleth'
True
```
```
>>> 'Cameroon' in fig.data[0]['hovertext']
True
```
failure_message: Ensuring that the `hover_name` argument is set correctly (make sure it's not set to the 'ISO' column, and make sure you didn't set the `hover_data` argument instead)
```
>>> fig.layout['coloraxis']['colorbar']['title']['text'] == 'Doses Per Person'
True
```
failure_message: Ensures the legend is set correctly
```
>>> 'covid vaccine doses per person' in fig.layout['title']['text'].lower()
True
```
failure_message: Ensures the title is set correctly
```
>>> colors = [c[1] for c in fig.layout['coloraxis']['colorscale']]
>>> default_colorscale = ['#0d0887', '#46039f', '#7201a8', '#9c179e', '#bd3786', '#d8576b', '#ed7953', '#fb9f3a', '#fdca26', '#f0f921']
>>> colors != default_colorscale
True
```
failure_message: Ensuring changes were made to the colorscale
```
>>> plotted_data = pd.Series(fig.data[0]['z'], index=fig.data[0]['locations'])
>>> np.isclose(plotted_data['USA'], 1.9496818625192982, atol=0.5)
True
```
failure_message: Plotted the correct data for the USA. If this fails, and your plot looks right, make sure that you only set the `locations` argument, and not the `locationmode` argument.


Tests for q6 (points: None)
```
>>> isinstance(israel, pd.DataFrame)
True
```
```
>>> israel.shape == (6937542, 3)
True
```


Tests for q7 (points: None)
```
>>> isinstance(test_stats_vax, np.ndarray) and isinstance(test_stats_sick, np.ndarray)
True
```
```
>>> len(test_stats_vax) == len(test_stats_sick) == 100
True
```
```
>>> missingness_type() in {1, 2, 3, 4}
True
```


Tests for q8 (points: None)
```
>>> isinstance(eff_overall, float)
True
```
```
>>> 0 <= eff_overall <= 1
True
```
```
>>> example_vax = pd.DataFrame({'Age': [15, 20, 25, 30, 35, 40], 
...                             'Vaccinated': [True, True, True, False, False, False], 
...                             'Severe Sickness': [True, False, False, False, True, True]})
>>> effectiveness(example_vax) == 0.5
True
```


Tests for q9 (points: None)
```
>>> isinstance(s_eff, pd.Series) 
True
```
```
>>> len(s_eff) == 10
True
```
```
>>> israel_raw = pd.read_csv(data_dir / 'israel-subset.csv')
>>> vax_subset = clean_israel_data(israel_raw).dropna()
>>> stratified_effectiveness(vax_subset).index[0] == '12-15'
True
```
```
>>> israel_raw = pd.read_csv(data_dir / 'israel-subset.csv')
>>> vax_subset = clean_israel_data(israel_raw).dropna()
>>> len(stratified_effectiveness(vax_subset)) == 10
True
```
