
# Project 4 Validation Notebook

You cannot write any code in this notebook – to actually work on the project, open `project.py` and `project.ipynb`.

The only purpose of this notebook is to give you a blank copy of `project.ipynb`,
so that you can go to **Kernel > Restart & Run All** and ensure that all public `grader.check` cells pass using just the code in your `project.py`.

**Before submitting Project 4, make sure that the call to `grader.check_all()` at the bottom of this notebook shows that all test cases passed!** 
If it doesn't, there's likely a function in `project.ipynb` that is not implemented correctly in `project.py`, or it could be that a function in `project.py` depends on an object (e.g. a DataFrame) that is not an argument to that function.
    

```python
# Initialize Otter
import otter
grader = otter.Notebook("project.ipynb")
```

```python
%load_ext autoreload
%autoreload 2
```

```python
import pandas as pd
import numpy as np
import os
import re
import requests
import time
```

```python
from project import *
```

```python
# don't change this cell, but do run it -- it is needed for the tests
beowulf = get_book('https://www.gutenberg.org/ebooks/16328.txt.utf-8')
```

```python
grader.check("q1")
```

```python
# don't change this cell, but do run it -- it is needed for the tests
import time
start = time.time()
shakes = tokenize(open('data/shakespeare.txt').read())
elapsed = time.time() - start
elapsed # Should be (much) under 5
```

```python
grader.check("q2")
```

```python
# don't change this cell, but do run it -- it is needed for the tests
tokens = tuple('one one two three one two four'.split())
unif = UniformLM(tokens)
unif.sample(100)
```

```python
grader.check("q3")
```

```python
# don't change this cell, but do run it -- it is needed for the tests
tokens = tuple('one one two three one two four'.split())
unigram = UnigramLM(tokens)
unigram.sample(100)
```

```python
grader.check("q4")
```

```python
# Uncomment and run – should take less than 10 seconds
# shakes_uniform = UniformLM(shakes)
# shakes_unigram = UnigramLM(shakes)
```

```python
# Uncomment and run
# shakes_uniform.sample(10)
```

```python
# Uncomment and run
# shakes_unigram.sample(10)
```

```python
# don't change this cell, but do run it -- it is needed for the tests
tokens = tuple('\x02 one two three one four \x03'.split())
bigrams = NGramLM(2, tokens)
out = bigrams.create_ngrams(tokens)
```

```python
# don't change this cell, but do run it -- it is needed for the tests
tokens = "\x02 Humpty Dumpty sat on a wall , Humpty Dumpty had a great fall . \x03 \x02 All the king ' s horses and all the king ' s men couldn ' t put Humpty together again . \x03".split()
tokens = tuple(tokens)
ngram = NGramLM(2, tokens)
out_5a1 = ngram.create_ngrams(tokens)
out_5b1 = ngram.mdl
out_5c1 = ngram
out_5d1 = ngram.sample(500) 
```

```python
grader.check("q5")
```

```python
# Uncomment and run
# shakes_bigram = NGramLM(2, shakes)
# shakes_bigram.sample(50)
```

```python
homer_tokens = tuple(open('data/homertokens.txt').read().split(' '))
```

```python
# NGramLM(5, homer_tokens)
```

```python
grader.check_all()
```


If you were able to go to **Kernel > Restart & Run All** and see all test cases pass above, and you've thoroughly tested your code yourself, you're ready to submit `project.py` to Gradescope!
    

Tests for q1 (points: None)
```
>>> isinstance(beowulf, str)
True
```
failure_message: check type to make sure the hidden tests will work
```
>>> beowulf[:20] == '\n\n\n\n\nProduced by Dav'
False
```
failure_message: make sure front matter not in string


Tests for q2 (points: None)
```
>>> 1000000 <= len(shakes) <= 1500000
True
```
failure_message: approx correct number of tokens
```
>>> elapsed <= 5
True
```
failure_message: shakespeare fast enough
```
>>> shakes[:3] == ['\x02', 'The', 'Complete'] and shakes[-3:] == ['William', 'Shakespeare', '\x03']
True
```
failure_message: check beginning and ending


Tests for q3 (points: None)
```
>>> (unif.mdl == 0.25).all()
True
```
failure_message: only one probability
```
>>> unif.mdl.shape[0] == 4
True
```
failure_message: number of tokens in mdl
```
>>> isinstance(unif.mdl, pd.Series)
True
```
failure_message: mdl correct type
```
>>> set(unif.mdl.index) == set('one two three four'.split())
True
```
failure_message: correct indices for mdl
```
>>> unif.probability(('five',)) == 0
True
```
failure_message: five not a token
```
>>> unif.probability(('one', 'two')) == 0.0625
True
```
failure_message: probability of given sequence appearing
```
>>> isinstance(unif.sample(1000), str)
True
```
failure_message: sample is string
```
>>> len(unif.sample(1000).split()) == 1000
True
```
failure_message: length of sample


Tests for q4 (points: None)
```
>>> unigram.mdl.nunique() == 3
True
```
failure_message: test mdl
```
>>> isinstance(unigram.mdl, pd.Series)
True
```
failure_message: mdl correct type
```
>>> set(unigram.mdl.index) == set('one two three four'.split())
True
```
failure_message: correct indices for mdl
```
>>> unigram.probability(('five',)) == 0
True
```
failure_message: five not a token
```
>>> isinstance(unigram.sample(1000), str)
True
```
failure_message: sample is string
```
>>> len(unigram.sample(1000).split()) == 1000
True
```
failure_message: sample length


Tests for q5 (points: None)
```
>>> isinstance(out[0], tuple)
True
```
failure_message: test create_ngrams
```
>>> out[0] == ('\x02', 'one') and out[2] == ('two', 'three')
True
```
failure_message: test create_ngrams
```
>>> set(bigrams.mdl.columns) == set('ngram n1gram prob'.split())
True
```
failure_message: test train
```
>>> bigrams.mdl.shape == (6, 3) and bigrams.mdl['prob'].min() == 0.5
True
```
failure_message: test train
```
>>> tokens = tuple('\x02 one two one three one two \x03'.split())
>>> bigrams2 = NGramLM(2, tokens)
>>> p = bigrams2.probability('two one three'.split())
>>> np.isclose(p, (1/4) * (1/2) * (1/3), 0.01)
True
```
failure_message: test probability
```
>>> samp = bigrams.sample(3)
>>> len(samp.split()) == 4  # don't count the initial START token.
True
```
failure_message: test sample
```
>>> samp = bigrams.sample(3)
>>> samp[:2] == '\x02 ' and  set(samp.split()) <= {'\x02', '\x03', 'one', 'two', 'three', 'four'}
True
```
failure_message: test sample
```
>>> 36 <= len(out_5a1) <= 40 and all([len(x) == 2 for x in out_5a1]) and out_5a1[0] == ('\x02', 'Humpty')
True
```
failure_message: test sample
```
>>> np.isclose(out_5b1.groupby('n1gram')['prob'].sum(), 1.0, 1.0).all() and (out_5b1['n1gram'].apply(len) == 1).all()
True
```
failure_message: test sample
```
>>> (out_5b1['ngram'].apply(len) == 2).all() and \
... np.isclose(out_5b1.loc[out_5b1.ngram == ('Humpty', 'Dumpty'), 'prob'].squeeze(), 0.666666, atol=0.1)
True
```
failure_message: test sample
```
>>> np.isclose(out_5c1.probability('Humpty Dumpty sat on a great fall'.split()), 0.01282051282051282, atol=0.1)
True
```
failure_message: test sample
```
>>> vc = pd.Series(out_5d1.split()).value_counts(normalize=True)
>>> {'\'', 'the', 'Humpty'}.intersection(set(vc.index[:6])) != {}
True
```
failure_message: most common sampled tokens
