
# Lab 5 Validation Notebook

You cannot write any code in this notebook – to actually work on the lab, open `lab.py` and `lab.ipynb`.

The only purpose of this notebook is to give you a blank copy of `lab.ipynb`,
so that you can go to **Kernel > Restart & Run All** and ensure that all public `grader.check` cells pass using just the code in your `lab.py`.

**Before submitting Lab 5, make sure that the call to `grader.check_all()` at the bottom of this notebook shows that all test cases passed!** 
If it doesn't, there's likely a function in `lab.ipynb` that is not implemented correctly in `lab.py`, or it could be that a function in `lab.py` depends on an object (e.g. a DataFrame) that is not an argument to that function.
    

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab.ipynb")
```

```python
%load_ext autoreload
%autoreload 2
```

```python
from lab import *
```

```python
import os
import pandas as pd
import numpy as np
from scipy import stats

import plotly.express as px
pd.options.plotting.backend = 'plotly'

import requests
import bs4
```

```python
# don't change this cell -- it is needed for the tests to work
out_q1 = after_purchase()
```

```python
grader.check("q1")
```

```python
# don't change this cell -- it is needed for the tests to work
out_q2 = multiple_choice()
```

```python
grader.check("q2")
```

```python
# Run your permutation tests in the Jupyter Notebook;
# put your final results in lab.py
payments_fp = os.path.join('data', 'payment.csv')
payments = pd.read_csv(payments_fp)
payments.head()
```

```python
# don't change this cell, but do run it -- it is needed for the tests to work
first_pval, first_result = first_round()
second_pval, second_result, second_result1 = second_round()
```

```python
grader.check("q3")
```

```python
# don't change this cell, but do run it -- it is needed for the tests to work
heights_fp = os.path.join('data', 'missing_heights.csv')
heights = pd.read_csv(heights_fp)
out_q4 = verify_child(heights.copy())
```

```python
grader.check("q4")
```

```python
out_q4.to_frame().rename(columns={0: 'p-value'}).plot(kind='barh', width=800, height=400)
```

```python
# don't change this cell, but do run it -- it is needed for the tests to work
heights_fp = os.path.join('data', 'missing_heights.csv')
new_heights = pd.read_csv(heights_fp)[['father', 'child_50']]
new_heights = new_heights.rename(columns={'child_50': 'child'})
out_q5 = cond_single_imputation(new_heights.copy())
```

```python
# don't change this cell, but do run it -- it is needed for the tests to work
heights_q5 = pd.read_csv('data/missing_heights.csv')
heights_q5['child'] = heights_q5['child_50']
inp_q5 = heights_q5
out_q5 = cond_single_imputation(inp_q5)
df_q5 = inp_q5.copy()
df_q5['imputed'] = out_q5
gp1_q5 = df_q5.groupby('father')['imputed'].mean()
gp2_q5 = df_q5.groupby('father')['child'].mean()
m_q5 = (pd.concat([gp1_q5, gp2_q5], axis=1)
     .dropna().diff(axis=1).abs().iloc[:, -1])
```

```python
grader.check("q5")
```

```python
# don't change this cell, but do run it -- it is needed for the tests to work
heights_fp = os.path.join('data', 'missing_heights.csv')
heights = pd.read_csv(heights_fp)
child = heights['child_50']
quantitative_distribution_out_q6 = quantitative_distribution(child.copy(), 100)
impute_height_quant_out_q6 = impute_height_quant(child.copy())
```

```python
grader.check("q6")
```

```python
# don't change this cell, but do run it -- it is needed for the tests to work
mc_answers, websites = answers()
```

```python
# don't change this cell, but do run it -- it is needed for the tests to work
from urllib.parse import urlparse
import urllib.robotparser

# This code checks the robots.txt file
def canFetch(url):
    if url[:4] != 'http':
        url = 'https://' + url
    parsed_uri = urlparse(url)

    domain = '{uri.scheme}://{uri.netloc}/'.format(uri=parsed_uri)

    rp = urllib.robotparser.RobotFileParser()
    rp.set_url(domain + "/robots.txt")
    try:
        rp.read()
        canFetchBool = rp.can_fetch("*", url)
    except:
        canFetchBool = None
    
    return canFetchBool

ans_q7, websites_q7 = answers()
canfetch_q7 = []
for site in websites_q7:
    try:
        canfetch_q7.append(canFetch(site))
    except:
        canfetch_q7.append(None)
```

```python
grader.check("q7")
```

```python
grader.check_all()
```


If you were able to go to **Kernel > Restart & Run All** and see all test cases pass above, and you've thoroughly tested your code yourself, you're ready to submit `lab.py` to Gradescope!
    

Tests for q1 (points: None)
```
>>> len(out_q1) == 5
True
```
failure_message: output length should be 5
```
>>> set(out_q1) <= set(['MD', 'MCAR', 'MAR', 'NMAR'])
True
```
failure_message: output contains answers other than the specified options
```
>>> out_q1[0] == 'NMAR'
True
```
failure_message: sub-question 1: reviewers are more likely to review when dissatified
```
>>> out_q1[0] in ['NMAR', 'MCAR']
True
```
failure_message: sub-question 1: partial; one column must be NMAR or MCAR
```
>>> out_q1[1] == 'MD'
True
```
failure_message: sub-question 2: not the optimal answer
```
>>> out_q1[1] in ['MD', 'MAR']
True
```
failure_message: sub-question 2: partial
```
>>> out_q1[2] == 'MAR'
True
```
failure_message: sub-question 3: not the optimal answer
```
>>> out_q1[2] in ['NMAR', 'MAR']
True
```
failure_message: sub-question 3: partial
```
>>> out_q1[3] in ['NMAR', 'MAR']
True
```
failure_message: sub-question 4: not the optimal answer, how does having a serial number affect missingness?
```
>>> out_q1[3] in ['NMAR', 'MAR', 'MCAR']
True
```
failure_message: sub-question 4: partial
```
>>> out_q1[4] == 'MAR'
True
```
failure_message: sub-question 5: not the optimal answer
```
>>> out_q1[4] in ['MAR', 'NMAR', 'MCAR']
True
```
failure_message: sub-question 5: partial


Tests for q2 (points: None)
```
>>> len(out_q2) == 5
True
```
failure_message: output length should be 5
```
>>> set(out_q2) <= set(['MD', 'MCAR', 'MAR', 'NMAR'])
True
```
failure_message: output contains answers other than the specified options
```
>>> out_q2[0] == 'MAR'
True
```
failure_message: sub-question 1
```
>>> out_q2[0] in ['MCAR', 'MAR', 'NMAR']
True
```
failure_message: sub-question 1: partial
```
>>> out_q2[1] in ['MAR', 'NMAR']
True
```
failure_message: sub-question 2
```
>>> out_q2[2] in ['MD', 'MAR']
True
```
failure_message: sub-question 3
```
>>> out_q2[3] == 'NMAR'
True
```
failure_message: sub-question 4
```
>>> out_q2[3] in ['NMAR', 'MCAR']
True
```
failure_message: sub-question 4: partial
```
>>> out_q2[4] == 'MCAR'
True
```
failure_message: sub-question 5
```
>>> out_q2[4] in ['MAR', 'NMAR', 'MCAR']
True
```
failure_message: sub-question 5: partial


Tests for q3 (points: None)
```
>>> isinstance(first_round(), list) and isinstance(second_round(), list)
True
```
```
>>> len(first_round()) == 2 and len(second_round()) == 3
True
```
```
>>> first_pval < 1
True
```
```
>>> first_result in ['NR', 'R']
True
```
```
>>> second_pval < 1
True
```
```
>>> second_result in ['NR', 'R']
True
```
```
>>> second_result1 in ['ND', 'D']
True
```
```
>>> np.isclose(first_pval, 0.16, atol=0.4)
True
```
failure_message: first_round p-value approximate
```
>>> first_result == 'NR'
True
```
failure_message: first_round result -- reject / do not reject
```
>>> second_result1 == 'D'
True
```
failure_message: part 3 multiple choice
```
>>> np.isclose(second_pval, 0.02, atol=0.2)
True
```
failure_message: second_round p-value approximate
```
>>> np.isclose(second_pval, 0.02, atol=0.1)
True
```
failure_message: second_round p-value approximate
```
>>> np.isclose(second_pval, 0.02, atol=0.02)
True
```
failure_message: second_round p-value
```
>>> second_result == 'R'
True
```
failure_message: second_round result -- reject / do not reject


Tests for q4 (points: None)
```
>>> out_q4['child_50'] < out_q4['child_95']
True
```
```
>>> out_q4['child_5'] > out_q4['child_50']
True
```
```
>>> out_q4['child_95'] > out_q4['child_5']
True
```
failure_message: 95 > 5
```
>>> out_q4['child_75'] > out_q4['child_25']
True
```
failure_message: 75 > 25
```
>>> np.isclose(out_q4['child_50'], 0, atol=0.02)
True
```
failure_message: pvalue of child_50
```
>>> np.isclose(out_q4['child_25'], 0.06, atol=0.2)
True
```
failure_message: pvalue of child_25
```
>>> np.isclose(out_q4['child_95'], 0.8, atol=0.2)
True
```
failure_message: pvalue of child_95


Tests for q5 (points: None)
```
>>> out_q5.isna().sum() == 0
True
```
```
>>> (new_heights['child'].std() - out_q5.std()) > 0.5
True
```
```
>>> np.isclose(out_q5.mean(), inp_q5['child'].mean(),atol=0.2)
True
```
failure_message: means are close
```
>>> np.all(np.isclose(m_q5, 0, atol=2))
True
```
failure_message: group-wise means are close: very approx
```
>>> np.isclose(m_q5.median(), 0.314, atol=0.1)
True
```
failure_message: group-wise means are close: approx
```
>>> np.isclose(m_q5.median(), 0.31439, atol=0.01)
True
```
failure_message: group-wise means are close


Tests for q6 (points: None)
```
>>> quantitative_distribution_out_q6.min() >= 56
True
```
```
>>> quantitative_distribution_out_q6.max() <= 79
True
```
```
>>> np.isclose(quantitative_distribution_out_q6.mean(), child.mean(), atol=1)
True
```
```
>>> np.isclose(quantitative_distribution_out_q6.std(), 3.5, atol=0.65)
True
```
```
>>> impute_height_quant_out_q6.isna().sum() == 0
True
```
```
>>> np.isclose(impute_height_quant_out_q6.mean(), child.mean(), atol=0.5)
True
```
```
>>> (set(quantitative_distribution_out_q6) - set(child.dropna())) != {}
True
```
failure_message: new values
```
>>> (set(impute_height_quant_out_q6) - set(child.dropna())) != {}
True
```
failure_message: new values


Tests for q7 (points: None)
```
>>> len(mc_answers) == 4
True
```
```
>>> len(websites) >= 2
True
```
```
>>> ans_q7[0] == 1
True
```
failure_message: multiple choice: 1
```
>>> ans_q7[1] == 2
True
```
failure_message: multiple choice: 2
```
>>> ans_q7[2] == 2
True
```
failure_message: multiple choice: 3
```
>>> ans_q7[3] == 1
True
```
failure_message: multiple choice: 4
```
>>> ('waymo.com' not in ''.join(websites)) and ('linkedin.com' not in ''.join(websites))
True
```
failure_message: no waymo or linkedin!
```
>>> True in canfetch_q7
True
```
failure_message: at least one that can scrape
```
>>> False in canfetch_q7
True
```
failure_message: at least one that can not scrape
