
# Lab 3 Validation Notebook

You cannot write any code in this notebook – to actually work on the lab, open `lab.py` and `lab.ipynb`.

The only purpose of this notebook is to give you a blank copy of `lab.ipynb`,
so that you can go to **Kernel > Restart & Run All** and ensure that all public `grader.check` cells pass using just the code in your `lab.py`.

**Before submitting Lab 3, make sure that the call to `grader.check_all()` at the bottom of this notebook shows that all test cases passed!** 
If it doesn't, there's likely a function in `lab.ipynb` that is not implemented correctly in `lab.py`, or it could be that a function in `lab.py` depends on an object (e.g. a DataFrame) that is not an argument to that function.
    

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab.ipynb")
```

```python
%load_ext autoreload
%autoreload 2
```

```python
from lab import *
```

```python
import os
import io
from pathlib import Path
import pandas as pd
import numpy as np
```

```python
# don't change this cell -- it is needed for the tests to work
fp = Path('data') /'universities_unified.csv'
df = pd.read_csv(fp)
cleaned = clean_universities(df)
info = university_info(cleaned)
```

```python
grader.check("q1")
```

```python
# do not edit this cell -- it is needed for the tests
fp = Path('data') / 'universities_unified.csv'
universities = pd.read_csv(fp)
cleaned = clean_universities(universities)
universities_out = std_scores_by_nation(cleaned)
su_and_spread_out = su_and_spread()
```

```python
grader.check("q2")
```

```python
# do not edit this cell -- it is needed for the tests
dirname = Path('data') / 'responses'
q3_out = read_linkedin_survey(dirname)
stats_out = com_stats(q3_out)
```

```python
grader.check("q3")
```

```python
# do not edit this cell -- it is needed for the tests
dirname = Path('data')/'extra-credit-surveys'
q4_out = read_student_surveys(dirname)
check_credit_out = check_credit(q4_out)
```

```python
grader.check("q4")
```

```python
# do not edit this cell -- it is needed for the tests
pets_fp = Path('data') / 'pets' / 'Pets.csv'
procedure_history_fp =  Path('data') / 'pets' / 'ProceduresHistory.csv'
owners_fp = Path('data') / 'pets' / 'Owners.csv'
procedure_detail_fp = Path('data') / 'pets' / 'ProceduresDetails.csv'
pets = pd.read_csv(pets_fp)
procedure_history = pd.read_csv(procedure_history_fp)
owners = pd.read_csv(owners_fp)
procedure_detail = pd.read_csv(procedure_detail_fp)

out_01 = most_popular_procedure(pets, procedure_history)
out_02 = pet_name_by_owner(owners, pets)
out_03 = total_cost_per_city(owners, pets, procedure_history, procedure_detail)
```

```python
grader.check("q5")
```

```python
# don't change this cell -- it is needed for the tests to work
fp = Path('data') / 'sales.csv'
sales = pd.read_csv(fp)
q6_average_seller_out = average_seller(sales)
q6_product_name_out = product_name(sales)
q6_product_count_out = count_product(sales)
q6_total_by_month_out = total_by_month(sales)
```

```python
grader.check("q6")
```

```python
grader.check_all()
```


If you were able to go to **Kernel > Restart & Run All** and see all test cases pass above, and you've thoroughly tested your code yourself, you're ready to submit `lab.py` to Gradescope!
    

Tests for q1 (points: None)
```
>>> cleaned.shape[0] == df.shape[0]
True
```
```
>>> cleaned['nation'].nunique() == 59
True
```
```
>>> len(info) == 4
True
```
```
>>> all([isinstance(x, y) for x, y in zip(info, [str, float, int, str])])
True
```
```
>>> (info[1] >= 0) & (info[1] <= 1)
True
```


Tests for q2 (points: None)
```
>>> universities_out.shape[0] == cleaned.shape[0]
True
```
```
>>> all(universities_out.columns == ['institution', 'nation', 'score'])
True
```
```
>>> np.all(abs(universities_out.select_dtypes(include='number').mean()) < 10**-7)  # standard units should average to 0!
True
```
```
>>> len(su_and_spread_out) == 2
True
```
```
>>> su_and_spread_out[0] in np.arange(1, 4)
True
```
```
>>> isinstance(su_and_spread_out[1], str)
True
```


Tests for q3 (points: None)
```
>>> isinstance(q3_out, pd.DataFrame) == True
True
```
```
>>> len(q3_out) == 5000
True
```
```
>>> try:
...     read_linkedin_survey('nonexistentfile')
... except FileNotFoundError:
...     print(True)
True
```
```
>>> len(stats_out) == 4
True
```
```
>>> isinstance(stats_out[0], float)
True
```
```
>>> isinstance(stats_out[1], int)
True
```
```
>>> isinstance(stats_out[2], str)
True
```


Tests for q4 (points: None)
```
>>> isinstance(q4_out, pd.DataFrame)
True
```
```
>>> q4_out.shape == (1000, 6)
True
```
```
>>> try:
...     read_student_surveys('nonexistentfile')
... except FileNotFoundError:
...     print(True)
True
```
```
>>> check_credit_out.shape == (1000, 2)
True
```
```
>>> check_credit_out['ec'].max() == 6
True
```


Tests for q5 (points: None)
```
>>> isinstance(out_01, str)
True
```
```
>>> len(out_02) == len(owners)
True
```
```
>>> 'Sarah' in out_02.index
True
```
```
>>> 'Cookie' in out_02.values
True
```
```
>>> set(out_03.index) <= set(owners['City'])
True
```


Tests for q6 (points: None)
```
>>> q6_average_seller_out.shape[0] == 3
True
```
```
>>> q6_average_seller_out["Average Sales"].sum() < 5000
True
```
```
>>> q6_product_name_out.size == 15
True
```
```
>>> q6_product_name_out["pen"].isnull().sum() == 0
True
```
```
>>> q6_product_count_out.size == 70
True
```
```
>>> q6_total_by_month_out.shape[1] == 5
True
```
```
>>> int(q6_product_name_out['hotel'].isnull().sum()) == 2
True
```
failure_message: null values: hotel row
```
>>> q6_product_count_out.shape == (10, 7)
True
```
failure_message: correct shape
```
>>> (q6_total_by_month_out > 2000).sum().sum() == 3
True
```
failure_message: number of entries strictly over 2000
```
>>> (q6_total_by_month_out == 0).sum().sum() == 39
True
```
failure_message: number of zero entries
