
# Lab 9 Validation Notebook

You cannot write any code in this notebook – to actually work on the lab, open `lab.py` and `lab.ipynb`.

The only purpose of this notebook is to give you a blank copy of `lab.ipynb`,
so that you can go to **Kernel > Restart & Run All** and ensure that all public `grader.check` cells pass using just the code in your `lab.py`.

**Before submitting Lab 9, make sure that the call to `grader.check_all()` at the bottom of this notebook shows that all test cases passed!** 
If it doesn't, there's likely a function in `lab.ipynb` that is not implemented correctly in `lab.py`, or it could be that a function in `lab.py` depends on an object (e.g. a DataFrame) that is not an argument to that function.
    

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab.ipynb")
```

```python
%load_ext autoreload
%autoreload 2
```

```python
from lab import *
```

```python
import pandas as pd
import numpy as np
import os
import plotly.express as px

from sklearn.linear_model import LinearRegression
from sklearn.preprocessing import FunctionTransformer
from sklearn.preprocessing import OneHotEncoder
from sklearn.pipeline import Pipeline
from sklearn.compose import ColumnTransformer

from pipeline_testing_util import get_transformers

from plotly.subplots import make_subplots
import plotly.graph_objects as go

import warnings
warnings.simplefilter('ignore')
```

```python
fp = os.path.join('data', 'toy.csv')
data = pd.read_csv(fp)
data.head()
```

```python
# don't change this cell, but do run it -- it is needed for the tests to work
q1_fp = os.path.join('data', 'toy.csv')
q1_data = pd.read_csv(q1_fp)
q1_pl, q1_preds = simple_pipeline(q1_data)
```

```python
grader.check("q1")
```

```python
# don't change this cell, but do run it -- it is needed for the tests to work
q2_fp = os.path.join('data', 'toy.csv')
q2_data = pd.read_csv(q2_fp)
q2_pl, q2_preds = multi_type_pipeline(q2_data)
```

```python
grader.check("q2")
```

```python
# The scatter plot referenced at the start of Question 3
# This is not needed to answer the question, but motivates why we are standardizing
px.scatter(data, x='c1', y='y', color='group')
```

```python
# don't change this cell, but do run it -- it is needed for the tests to work
# test fit 
q3_test_fit_cols = {'g': ['A', 'A', 'B', 'B'], 'c1': [1, 2, 2, 2], 'c2': [3, 1, 2, 0]}
q3_test_fit_X = pd.DataFrame(q3_test_fit_cols)
q3_test_fit_std = StdScalerByGroup().fit(q3_test_fit_X)

# test transform
q3_test_transform_cols = {'g': ['A', 'A', 'B', 'B'], 'c1': [1, 2, 3, 4], 'c2': [1, 2, 3, 4]}
q3_test_transform_X = pd.DataFrame(q3_test_transform_cols)
q3_test_transform_std = StdScalerByGroup().fit(q3_test_transform_X)
q3_test_transform_out = q3_test_transform_std.transform(q3_test_transform_X)
```

```python
# don't change this cell, but do run it -- it is needed for the tests to work
q3_fit_data = pd.read_csv('data/toy.csv')

N = 2*10**6
a = np.random.choice(['A', 'B'], size=(N,1)).astype('object')
b = np.random.multivariate_normal([1, 2], [[1, 0],[0, 100]], size=N)
arr = np.hstack([a, b])
q3_transform_data = pd.DataFrame(arr)
q3_transform_data[1] = q3_transform_data[1].astype(float)
q3_transform_data[2] = q3_transform_data[2].astype(float)
```

```python
grader.check("q3")
```

```python
grader.check("q4")
```

```python
# Use `galton` to test your work.
galton = pd.read_csv('data/galton.csv')
galton.head()
```

```python
# don't change this cell, but do run it -- it is needed for the tests
galton_test = pd.read_csv('data/galton.csv')
out_tree_test = tree_reg_perf(galton_test)
out_knn_test = knn_reg_perf(galton_test)
```

```python
grader.check("q5")
```

```python
np.random.seed(9) # For reproducibility

tree = tree_reg_perf(galton)
knn = knn_reg_perf(galton)
hyp = np.arange(1, 21)

fig = make_subplots(rows=1, cols=2, subplot_titles=('Error vs. Tree Depth<br>for Decision Tree Regressor',
                                                    'Error vs. k (# Neighbors)<br>for k-NN Regressor'))

fig.add_trace(
    go.Scatter(x=hyp, y=tree.iloc[:, 0], line=dict(color='blue'), name='Training Error'),
    row=1, col=1).add_trace(go.Scatter(x=hyp, y=tree.iloc[:, 1], line=dict(color='red'), name='Testing Error'), 
                            row=1, col=1)

fig.add_trace(
    go.Scatter(x=hyp, y=knn.iloc[:, 0], line=dict(color='blue'), name='Training Error', showlegend=False),
    row=1, col=2).add_trace(go.Scatter(x=hyp, y=knn.iloc[:, 1], line=dict(color='red'),  name='Testing Error',
                                       showlegend=False), row=1, col=2)

fig.update_layout(height=450, width=975)
fig.update_xaxes(title_text='Tree Depth', row=1, col=1, tickvals=np.arange(1,21,2))
fig.update_xaxes(title_text='k (# Neighbors)', row=1, col=2, tickvals=np.arange(1,21,2))

fig.show()
```

```python
# Experiment using `titanic` below – remember, this is only your training data
titanic = pd.read_csv('data/titanic.csv')
titanic.head()
```

```python
from sklearn.model_selection import train_test_split
X = titanic.drop(columns = 'Survived')
y = titanic['Survived']
X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.35)
pre = titanic_model(titanic)
(pre.predict(X_test)==y_test).mean()
```

```python
# don't change this cell, but do run it -- it is needed for the tests
q6_data_test = pd.read_csv('data/titanic.csv')
pl_test = titanic_model(q6_data_test)
```

```python
grader.check("q6")
```

```python
grader.check_all()
```


If you were able to go to **Kernel > Restart & Run All** and see all test cases pass above, and you've thoroughly tested your code yourself, you're ready to submit `lab.py` to Gradescope!
    

Tests for q1 (points: None)
```
>>> isinstance(q1_pl, Pipeline)
True
```
```
>>> isinstance(q1_pl.steps[-1][1], LinearRegression)
True
```
```
>>> isinstance(q1_pl.steps[0][1], FunctionTransformer)
True
```
```
>>> q1_preds.shape[0] == q1_data.shape[0]
True
```


Tests for q2 (points: None)
```
>>> isinstance(q2_pl, Pipeline)
True
```
```
>>> isinstance(q2_pl.steps[-1][1], LinearRegression)
True
```
```
>>> isinstance(q2_pl.steps[0][1], ColumnTransformer)
True
```
```
>>> q2_data.shape[0] == q2_preds.shape[0]
True
```


Tests for q3 (points: None)
```
>>> q3_test_fit_std.grps_ is not None
True
```
```
>>> q3_test_transform_out.shape == (4, 2)
True
```
```
>>> np.isclose(q3_test_transform_out.abs(), 0.707107, atol=0.001).all().all()
True
```


Tests for q4 (points: None)
```
>>> out = eval_toy_model()
>>> len(out) == 3
True
```
```
>>> out = eval_toy_model()
>>> np.all([len(t) == 2 for t in out])
True
```


Tests for q5 (points: None)
```
>>> (out_tree_test.columns.tolist() == ['train_err', 'test_err']) == True
True
```
```
>>> (out_tree_test['train_err'].iloc[-1] < out_tree_test['test_err'].iloc[-1]) == True
True
```
```
>>> (out_knn_test.columns.tolist() == ['train_err', 'test_err']) == True
True
```


Tests for q6 (points: None)
```
>>> isinstance(pl_test, Pipeline)
True
```
```
>>> isinstance(pl_test.steps[-1][-1], BaseEstimator)
True
```
```
>>> preds_test = pl_test.predict(q6_data_test.drop('Survived', axis=1))
>>> ((preds_test == 0) | (preds_test == 1)).all()
True
```
