
# Lab 7 Validation Notebook

You cannot write any code in this notebook – to actually work on the lab, open `lab.py` and `lab.ipynb`.

The only purpose of this notebook is to give you a blank copy of `lab.ipynb`,
so that you can go to **Kernel > Restart & Run All** and ensure that all public `grader.check` cells pass using just the code in your `lab.py`.

**Before submitting Lab 7, make sure that the call to `grader.check_all()` at the bottom of this notebook shows that all test cases passed!** 
If it doesn't, there's likely a function in `lab.ipynb` that is not implemented correctly in `lab.py`, or it could be that a function in `lab.py` depends on an object (e.g. a DataFrame) that is not an argument to that function.
    

```python
# Initialize Otter
import otter
grader = otter.Notebook("lab.ipynb")
```

```python
%load_ext autoreload
%autoreload 2
```

```python
from lab import *
```

```python
import pandas as pd
import numpy as np
import os
from pathlib import Path
import re
```

```python
grader.check("q1")
```

```python
# experiment with extract_personal using the file s below
fp = Path('data') / 'messy.txt'
s = open(fp, encoding='utf8').read()
```

```python
# don't change this cell, but do run it -- it is needed for the tests
test_fp = Path('data') / 'messy.txt'
test_s = open(test_fp, encoding='utf8').read()
emails, ssn, bitcoin, addresses = extract_personal(test_s)
```

```python
grader.check("q2")
```

```python
# experiment with tfidf_data using reviews_ser and review below 
fp = Path('data') / 'reviews.txt'
reviews_ser = pd.read_csv(fp, header=None).squeeze("columns")
review = open(Path('data') / 'reviews.txt', encoding='utf8').read().strip()
```

```python
# don't change this cell, but do run it -- it is needed for the tests
fp = os.path.join('data', 'reviews.txt')
reviews_ser = pd.read_csv(fp, header=None).squeeze("columns")
review = open(os.path.join('data', 'review.txt'), encoding='utf8').read().strip()
q3_tfidf = tfidf_data(reviews_ser, review)

try:
    q3_rel = relevant_word(q3_tfidf)
except:
    q3_rel = None
```

```python
grader.check("q3")
```

```python
# The public tests don't test your work on the `ira` data,
# but the hidden tests do.
# So, make sure to thoroughly test your work yourself!
fp = os.path.join('data', 'ira.csv')
ira = pd.read_csv(fp, names=['id', 'name', 'date', 'text'])
```

```python
grader.check("q4")
```

```python
# The doctests/public tests don't test your work on the `ira` data,
# but the hidden tests do.
# So, make sure to thoroughly test your work yourself!
fp = os.path.join('data', 'ira.csv')
ira = pd.read_csv(fp, names=['id', 'name', 'date', 'text'])
```

```python
# don't change this cell, but do run it -- it is needed for the tests
# (yes, we know it says "hidden" – there are still truly hidden tests in this question)
fp_hidden = 'data/ira_test.csv'
ira_hidden = pd.read_csv(fp_hidden, header=None)
text_hidden = ira_hidden.iloc[:, -1:]
text_hidden.columns = ['text']

test_hidden = create_features(text_hidden)
```

```python
grader.check("q5")
```

```python
grader.check_all()
```


If you were able to go to **Kernel > Restart & Run All** and see all test cases pass above, and you've thoroughly tested your code yourself, you're ready to submit `lab.py` to Gradescope!
    

Tests for q1 (points: None)
```
>>> import doctest
>>> doctest.run_docstring_examples(match_1, {'match_1': match_1})
>>> doctest.run_docstring_examples(match_2, {'match_2': match_2})
>>> doctest.run_docstring_examples(match_3, {'match_3': match_3})
>>> doctest.run_docstring_examples(match_4, {'match_4': match_4})
>>> doctest.run_docstring_examples(match_5, {'match_5': match_5})
>>> doctest.run_docstring_examples(match_6, {'match_6': match_6})
>>> doctest.run_docstring_examples(match_7, {'match_7': match_7})
>>> doctest.run_docstring_examples(match_8, {'match_8': match_8})
>>> doctest.run_docstring_examples(match_9, {'match_9': match_9})
>>> doctest.run_docstring_examples(match_10, {'match_10': match_10})
```
failure_message: doctests


Tests for q2 (points: None)
```
>>> emails[0] == 'test@test.com'
False
```
```
>>> ssn[0] == '423-01-9575'
False
```
```
>>> bitcoin[0] == '1BvBMSEYstWetqTFn5Au4m4GFg7xJaNVN2'
False
```
```
>>> addresses[0] == '530 High Street'
False
```


Tests for q3 (points: None)
```
>>> q3_tfidf['cnt'].sum() == 85
True
```
```
>>> 'before' in q3_tfidf.index
True
```
```
>>> q3_tfidf.loc['a', 'cnt'] == 4
True
```
```
>>> np.isclose(q3_tfidf.loc['phone', 'idf'], 0.820703, atol=0.5)
True
```
```
>>> q3_rel in q3_tfidf.index
True
```


Tests for q4 (points: None)
```
>>> testdata = [['RT @DSC80: Text-cleaning is cool! #NLP https://t.co/xsfdw88d #NLP1 #NLP1']]
>>> test = pd.DataFrame(testdata, columns=['text'])['text']
>>> out = hashtag_list(test)
>>> out.iloc[0] == ['NLP', 'NLP1', 'NLP1']
True
```
```
>>> testdata = [['RT @DSC80: Text-cleaning is cool! #NLP https://t.co/xsfdw88d #NLP1 #NLP1']]
>>> test = hashtag_list(pd.DataFrame(testdata, columns=['text'])['text'])
>>> most_common_hashtag(test).iloc[0] == 'NLP1'
True
```
```
>>> out = pd.Series([[2], [2], [3, 2, 3]])
>>> (most_common_hashtag(out) == 2).all()
True
```


Tests for q5 (points: None)
```
>>> testdata = [['RT @DSC80: Text-cleaning is cool! #NLP https://t.co/xsfdw88d #NLP1 #NLP1']]
>>> test = pd.DataFrame(testdata, columns=['text'])
>>> out = create_features(test)
>>> anscols = ['text', 'num_hashtags', 'mc_hashtags', 'num_tags', 'num_links', 'is_retweet']
>>> ansdata = [['text cleaning is cool', 3, 'NLP1', 1, 1, True]]
>>> ans = pd.DataFrame(ansdata, columns=anscols)
>>> (out == ans).all().all()
True
```
```
>>> test_hidden.columns.tolist() == ['text', 'num_hashtags', 'mc_hashtags', 'num_tags', 'num_links', 'is_retweet']
True
```
```
>>> 400 <= (test_hidden['num_tags'] == 0).sum() <= 450
True
```
```
>>> 450 <= (test_hidden['num_links'] == 0).sum() <= 550
True
```
